Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Liuda Serohina"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    #stores the tuple with start index, end index and value of max subarray
    result = []
    #stores values of the maximum subarray
    max_value = []
    
    #search from left to right
    for left_i in range(len(A)):
        for right_i in range(len(A)):
            if right_i > left_i:
                #add start, end and sum to result list
                result.append((left_i, right_i, sum(A[left_i:right_i+1])))
                #add the maximum value of the subarray to empty list
                max_value.append(sum(A[left_i:right_i+1]))

    #outputs a list with start value, end value and maximum value            
    return result[max_value.index(max(max_value))]

A = [-2, -3, 4, -1, -2, 1]
print(max_subarray(A))
```

Output:
```
(2, 3, 3)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. With the output from ```max_subarray(A)```, do you need to rerun ```max_subarray()``` for the new array again to find the subarray we need? Explain why, or why not, in as much detail as possible, using, if necessary, a specific example to illustrate your reasoning. The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
You don't have to rerun the entire maximum subarray because all the possible sums within it have already been computed. Thus, if you're adding a new element and assuming that it might change the maximum sum, you would need to perform comparisons from the start index (in our case 2) till the end (index 6). This summation would be the only one that wasn't performed yet. 

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
True, because the maximum sum of array B will only be greater if the last element of the array is 8. Otherwise, their maximum sums are the same. 

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    #if x has only 1 element, return it
    if len(x)==1:
        return x[0]
    
    #storage variable for subarray sums
    total_sum = 0
    
    #iterate thru the array in reverse order
    for i in range(len(x)-1,-1, -1):
        #add x[i] to total sum
        total_sum += x[i]
        #assign the total sum to the value of maximum sum if it is greater than maximum sum
        if total_sum > mx:
            mx = total_sum
    return mx
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    #setting the first element as the current maximum sum
    maximum_sum = A[0]
    
    #iterate for the whole length of array
    for x in range(0, len(A)):
        #calculate the max sum of subarrays using incremental_max_subarray
        maximum_sum = incremental_max_subarray(A[:x+1], maximum_sum)
        
    return float(maximum_sum)
    
max_subarray(B)
max_subarray(A)
```

Output:
```
4.0
```

```python
#test cases: aren't there 2 maximum subarrays here? We ge the value of 10(3+7)
#because we iterate from 0 to len(A). Is this correct or should the code output 2 maximum values?
C = [-5,3,7,-2,-9, 10]
assert(max_subarray(C) == 10)
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
I don't think so. The whole max_subarray would have the time complexity of Theta(n^2) because of the 2 for loops (one in max_subarray and another one in incremental_max_subarray). The time complexity of divide-and-conquer is Theta(n log n). Which means that as n tends to infinity, the time taken buy divide-and-conquer will be smaller than the time taken by max_subarray algorithm.