Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **50 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Liuda Serohina"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Analyzing elementary sorting algorithms



## Question 1 [time estimate: 10 minutes]
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
I think that line 7 means that we're decreasing the number of unsorted elements in the array. It is necessary to add this line since it would mean that both conditions (i>0 and A[i]>key) are satisfied, thus the element gets sorted. 

## Question 2 [time estimate: 15 minutes]
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertion_sort(A):
    for j in range(2, len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertion_sort(A):
    for j in range(1, len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
    #raise NotImplementedError()
```

```python
#test
A=[12,3,44]
insertion_sort(A)

for i in range(len(A)): 
    print ("%d" %A[i]) 
```

Output:
```
3
12
44
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertion_sort([0]) == [0])
assert(insertion_sort([-1,1]) == [-1,1])
assert(insertion_sort([1,-1]) == [-1,1])
assert(insertion_sort([1,6,3,6]) == [1,3,6,6])
assert(insertion_sort([7, 3, 5, 4, 6, 4]) == [3, 4, 4, 5, 6, 7])
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
I am not sure if this is correct, but by changing the inputs to the range function from (2, len(A)) to (1, len(A)), the error stopped showing up. I checked the code on a test array and it succeeded in sorting it. Nonetheless, the code cell with assert does not work. 

## Question 3 [time estimate: 15 minutes]
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertion_sort`,  your new Python function, `bubble_sort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubble_sort(A):
    n=len(A)
    
    #go through all the elements in an array
    for i in range(n-1): 
        #go through the array from 0 to n-i-1
        for j in range(0, n-i-1):
            #if the element you found is greater then the following one, swap them
            if A[j] > A[j+1]:
                A[j], A[j+1] = A[j+1], A[j]
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

```python
#test 1: regular check
A=[12,3,44]
bubble_sort(A)

for i in range(len(A)): 
    print ("%d" %A[i]) 
```

Output:
```
3
12
44
```

```python
#test 2: checking negative numbers
A=[-7,-1,-12]
bubble_sort(A)

for i in range(len(A)): 
    print ("%d" %A[i]) 
```

Output:
```
-12
-7
-1
```

```python
#test 3: checking with 0
A=[-7,2,0]
bubble_sort(A)

for i in range(len(A)): 
    print ("%d" %A[i]) 
```

Output:
```
-7
0
2
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
The code passing these three tests does not necessarily grant its correctness and there are multitudes of tests one can perform (checking with zeros, negative numbers, some other edge cases, numbers with decimals, etc) to gain more confidence. I suppose the program to be bug-free since Python doesn't throw any errors, although this does not mean that something might still be wrong and the code would execute the command with the embedded mistake.

## Question 4 [time estimate: 10 minutes]
The following pseudocode corresponds to a third sorting algorithm covered in this lesson, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free by providing at least three test cases. `selection_sort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selection_sort(A):
    n=len(A)
    
    #go through all the elements in the array
    for i in range(n):
        
        #search for the minimum value in the unsorted partition
        min_ind = i
        for j in range(i+1, n):
            #update the minimum number if the condition holds true
            if A[j] < A[min_ind]:
                min_ind = j
        #swap the new minimum number with the first element
        A[i], A[min_ind] = A[min_ind], A[i]
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

```python
A=[12,3,44]
selection_sort(A)

for i in range(len(A)): 
    print ("%d" %A[i]) 
```

Output:
```
3
12
44
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**