Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **63 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Liuda Serohina"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Computational Applications of Randomized Algorithms

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson Heaps and Heapsort, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.

## Question 1 [time estimate: 5 minutes]
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    """
    This function returns the median element
    """
    # Balancing min and max heaps by adding elements from one to the other
    if len(minh) > len(maxh)+1:
        heapq.heappush(maxh,-heapq.heappop(minh))
    if len(maxh) > len(minh)+1:
        heapq.heappush(minh,-heapq.heappop(maxh)) 
    # Median is a root of max heap if the max heap has more elements
    if len(minh) > len(maxh): 
        med = minh[0]
    # Median is a root of min heap if the min heap has more elements
    if len(minh) < len(maxh):
        med = maxh[0]
    # If the min and max heaps are the same, get the median using the following formula
    if len(minh) == len(maxh):
        med = ((-maxh[0])+minh[0])/2

    return median 
```

## Question 2 [time estimate: 10 minutes]
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq

def add_to_median_heap(minh, maxh, elem):    
    
    # Push an element into the max heap    
    if len(maxh) == 0:
        heapq.heappush(maxh, -elem)
    # When the new element is less than the median
    elif elem <= median(minh, maxh):
        # Add it to the min heap
        heapq.heappush(minh, elem)
    else: 
        # Otherwise, add it to the max heap
        heapq.heappush(maxh, -elem)
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
<function median at 0x7ff8d7d98c20>
<ipython-input-12-2a7cc854325b> in <module>
      2 maxh = []
      3 for a in range(1,100,2):
----> 4     add_to_median_heap(minh, maxh, a)
      5     print(median(minh, maxh))

<ipython-input-10-59d2087382c0> in add_to_median_heap(minh, maxh, elem)
      4     if len(maxh) == 0:
      5         heapq.heappush(maxh, -elem)
----> 6     elif elem <= median(minh, maxh):
      7         heapq.heappush(minh, elem)
      8     else:

TypeError: '<=' not supported between instances of 'int' and 'function'
```

## Question 4 [time estimate: 3 minutes]
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
The complexity of add_to_median_heap depends on the median function. The worst case complexity would happend when the heaps are extremely unbalanced. Then, the heappush would have the time complexity if O(log n). Thus, for n elements it becomes O(n log n). 

## Question 5 [time estimate: 3 minutes] 
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
For the median function, the worst case happens when there's a need to apply heappush. Thus, it is O(log n). 

## Question 6 [time estimate: 4 minutes]

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
It seems that this method is as good as other sorting algorithms since it has O(n log n). Something that's good about this algorithm is that you don't need to sort the list repeatedly when adding new elements. Using min and max heaps makes the structure of the algorithm fairly clear. 

## [Optional] Question 7 [time estimate: 15 minutes]

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1 [time estimate: 10 minutes]

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def partition(A):
    """
    This function partitions the input list.
    """
    x = A[-1]
    i = -1
    for j in range(len(A)):
        if A[j] <= x:
            i = i+1
            A[i], A[j] = A[j], A[i]
    A[i+1], A[-1] = A[-1], A[i+1]
    return i+1
```

```python
def qselect(lst,k):
    if len(lst) == 1:
        return lst[0]
    else:
        partition_idx = partition(lst)
        if partition_idx == k:
            return(lst[partition_idx])
        elif partition_idx > k:
            return qselect(lst[:partition_idx],k)
        else:
            k = k - partition_idx - 1
            return qselect(lst[partition_idx+1:],k)
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

Output:
```
<ipython-input-30-47dea53caced> in qselect(lst, k)
      3         return lst[0]
      4     else:
----> 5         partition_idx = partition(lst)
      6         if partition_idx == k:
      7             return(lst[partition_idx])

<ipython-input-29-f867756fce5d> in partition(A)
      9             i = i+1
     10             A[i], A[j] = A[j], A[i]
---> 11     A[i+1], A[-1] = A[-1], A[i+1]
     12     return i+1

IndexError: list index out of range
```

## Question 2 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.

```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))
```

Output:
```
<ipython-input-30-47dea53caced> in qselect(lst, k)
      3         return lst[0]
      4     else:
----> 5         partition_idx = partition(lst)
      6         if partition_idx == k:
      7             return(lst[partition_idx])

<ipython-input-29-f867756fce5d> in partition(A)
      9             i = i+1
     10             A[i], A[j] = A[j], A[i]
---> 11     A[i+1], A[-1] = A[-1], A[i+1]
     12     return i+1

IndexError: list index out of range
```

## Question 3 [time estimate: 5 minutes]

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
Since this algorithm is similar to quicksort, we can start with quicksort's recurrence equation: 2T(n/2) + Theta(n). Since we perform recursion only of one side of the list, we can modify it to be: T(n/2) + Theta(n).

## Question 4 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
The best case for this would be an equal split. I still don't really know how to solve thse since we didn't dedicate time to practicing this in class. Start by rewriting in the appropriate form T(n) = T(n/2) + Theta(n). I think as we keep decomposing the recursion tree into levels, it will boil down to n*1, thus O(n).

## Question 5 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
The worst case would mean a very unbalanced split. T(n) = T(n-1) + Theta(n). If we extrapolate from the analogy with quicksort, the worst case is also O(n^2). 